Stumpy forms are the predominant transmissible forms of Trypanosoma brucei

Author Response

The following is the authors’ response to the original reviews.

Public Reviews:

Reviewer #1:

Summary:

Ngoune et al. present compelling evidence that Slender cells are challenged to infect tsetse flies. They explore the experimental context of a recent important paper in the field, Schuster et al., that presents evidence suggesting the proliferative Slender bloodstream T. brucei can infect juvenile tsetse flies. Schuster et al. were disruptive to the widely accepted paradigm that the Stumpy bloodstream-form is solely responsible for tsetse infection and T. brucei transmission potential. Evidence presented here shows that in all cases, Stumpy form parasites are exponentially more capable of infecting tsetse flies. They further show that Slender cells do not infect mature flies.

However, they raise questions of immature tsetse immunological potential and field transmission potential that their experiments do not address. Specifically, they do not show that teneral tsetse flies are immunocompromised, that tsetse flies must be immunocompromised for Slender infection nor that younger teneral tsetse infection is not pertinent to field transmission.

Strengths:

Experimental Design is precise and elegant, outcomes are convincing. Discussion is compelling and important to the field. This is a timely piece that adds important data to a critical discussion of host: parasite interactions, of relevance to all parasite transmission.

Thank you

Weaknesses:

As above, the authors dispute the biological relevance of teneral tsetse infection in the wild, without offering evidence to the contrary. Statements need to be softened for claims regarding immunological competence or relevance to field transmission.

We have modified the revised version to soften these claims (l.156 and l.159). Please, note that the limited immunocompetence of teneral flies has been extensively studied by the labs of S. Aksoy at Yale and M. Lehane at Liverpool. In the discussion, we provide key references from these two labs 18-21. Our comment on the relevance to field transmission is simply based on field observations of the fly biology.

Reviewer #2:

Summary:

Contrary to findings recently reported by Schuster S et al., this short paper shows evidence that the stumpy form of T. brucei is probably the most pre-adapted form to progress with the life cycle of this parasite in the tsetse vector.

Strengths:

One of the most important pieces of experimental evidence is that they conduct all fly infection experiments in the absence of metabolites like GlcNAc or S-glutathione; by doing so, the infection rates in flies infected with slender trypanosomes seem very low or non-existent. This, on its own, is a piece of important experimental evidence that the Schuster S et al findings may need to be revisited.

Thank you

Weaknesses:

I consider that the authors should have included their own experiments demonstrating that the addition of these chemicals enhances the infection rates in flies receiving bloodmeals containing slender trypanosomes.

The main purpose of this study is to assess the intrinsic infectivity of SL Vs. ST in teneral Vs. adult flies, not to reproduce the results obtained by Schuster et al.. We think that the suggested experiment is not necessary as L-Glutathion is well-known to enhance infection rates by reducing the fly immune response efficiency (Ref 24). Most of the experimental infections with procyclic or ST forms (even at low densities) published by our lab and others, especially for studying parasite stages in the salivary glands, were actually performed by complementing the infective meal with L-Glutathion for this reason.

Reviewer #3:

The dogma in the Trypanosome field is that transmission by Tsetse flies is ensured by stumpy forms. This has been recently challenged by the Engstler lab (Schuster et al.), which showed that slender forms can also be transmitted by teneral flies. In this work, the authors aimed to test whether transmission by slender forms is possible and frequent.

For this, the authors repeated Tsetse transmission experiments but with some key critical differences relative to Schuster et al. First, they infected teneral and adult flies. Second, their infective meals lacked two components (N-acetylglucosamine and glutathione), which could have boosted the infection rates in the Schuster et al. work. In these conditions, the authors observed that most stumpy form infections with teneral and adult flies were successful while only 1 out of 24 slender-form infections was successful. Adult flies showed a lower infection rate, which is probably because their immune system is more developed.

Given that in Tsetse-infested areas most transmission is likely ensured by adult flies, the authors conclude that the parasite stage that will have a significant epidemiologic impact on transmission is the stumpy form.

Strengths:

• This work tackles an important question in the field.

• The Rotureau laboratory has well-known expertise in Tsetse fly transmission experiments.

• Experimental setup is robust and data is solid.

• The paper is concise and clearly written.

Thank you

Weaknesses:

• The reason(s) for why this work has lower infection rates with slender forms than Schuster et al. remain unknown. The authors suggested it could be because of the absence of N-acetylglucosamine and/or glutathione, but this was not formally tested. Could another source of variation be the clone of EATRO1125 AnTat1.1 (Paris versus Munich origin)? To reduce the workload, such additional experiments could be done with just one dose of parasites.

Differences between the strain clones, the cell culture conditions and/or the fly colony maintenance conditions could indeed explain the differences in infection rates observed in the two studies. However, the main purpose of this study is to assess the intrinsic infectivity of SL Vs. ST in teneral Vs. adult flies. Our study was designed to stand alone for providing a clear answer to this question, not to reproduce the results obtained by Schuster et al.. Hence, we don’t think that any additional experiments are required here.

• The characterization of what is slender and stumpy is critical. The authors used PAD1 protein expression as the sole reporter. While this is a robust assay to confirm stumpy, an analysis of the cell cycle would have been helpful to confirm that slender forms have not initiated differentiation (Larcombe S et al. 2023, preprint).

In this study, ST are indeed defined by their general morphology and by the expression of PAD1 proteins at the cell membrane as assessed by IFA. This is the simplest and most accurate ST proxy accessible by IFA. We do not think that monitoring in more details the cell cycle would provide key information here. If some SL forms had initiated differentiation in our experiments, then, the low infection rates observed with SL would have reinforced the fact that mostly mature PAD1+ ST are infectious for flies .

• Statistical analysis is missing. Is the difference between adult and teneral infections statistically significant?

An ANOVA statistical analysis was performed and a dedicated section was added to the revised version.

For all conditions, MG infection rate comparisons between adult and teneral flies were statistically significant.

Recommenda8ons for the authors:

Reviewer #1:

While some perceived outcomes pertaining to immunological competence and transmission relevance of teneral flies are overstated, the overall tone of the paper is inappropriately apologe7c. The authors obviously don't want to offend their colleagues but the current wri7ng style obscures meaning, making the paper a bit 'flowery' and difficult to read.

Ngoune et al. have important outcomes that need to be stated more directly.

Words such as 'unequivocally' are not appropriate to Schuster et al's outcomes. As your study shows, their findings are experimentally based, with inherent caveats, and are therefore sugges7ve, not demonstrated or proven.

The word 'unequivocally' has been removed from the revision.

Reviewer #3:

The Engstler lab cul7vates AntTaT1.1 in methylcellulose (Munich clone, if I am not mistaken). The Rotureau lab uses the Paris AntTaT1.1 clone and uses no methylcellulose. Given that methylcellulose helps stumpy forma7on, it seems important to show that the results of this paper are reproducible with the Munich clone grown in the presence of methylcellulose.

Differences between the strain clones and culture conditions could indeed explain the differences in infection rates observed in the two studies. However, the main purpose of this study is to assess the intrinsic infectivity of SL Vs. ST in teneral Vs. adult flies. Our study was designed to stand alone for providing a clear answer to this question, not to reproduce the results obtained by Schuster et al.. Hence, we don’t think that any additional experiments are required here.

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Recently, Schuster et al. demonstrated that bloodstream slender forms of African trypanosomes are readily transmissible to young tsetse flies where they can complete their complex life cycle (1). In their experimental conditions, a single slender parasite was sufficient for productive infection. Here, we compared the infectivity of slender and stumpy bloodstream forms in adult flies with a mature immune system, and without using any chemical compounds that would alter the insect immune response and/or promote the infection. After ingestion of slender forms, infected flies were observed only in one out of 24 batches of non-immunocompetent teneral flies and with a high number of parasites. In contrast, infected flies were detected in 75% (18/24) of the batches infected with stumpy parasites, and as few as 10 stumpy parasites produced mature infections in immune adult flies. We discuss that, although Schuster et al. have unequivocally demonstrated the intrinsic capacity of slender form trypanosomes to infect young and naive tsetse flies, highlighting the remarkable plasticity and adaptability of these protists, this phenomenon is unlikely to significantly contribute to the epidemiology of African trypanosomiases. According to both experimental and field observations, stumpy forms appear to be the most adapted forms for African trypanosome transmission from the mammalian host to the tsetse fly vector in natural conditions.

Revised version: Recently, Schuster et al. demonstrated that bloodstream slender forms of African trypanosomes are readily transmissible to young tsetse flies where they can complete their complex life cycle (1). In their experimental conditions, a single slender parasite was sufficient for productive infection. Here, we compared the infectivity of slender and stumpy bloodstream forms in adult flies with a mature immune system, and without using any chemical compounds that would alter the insect immune response and/or promote the infection. After ingestion of slender forms, infected flies were observed only in one out of 24 batches of non-immunocompetent teneral flies and with a high number of parasites. In contrast, infected flies were detected in 75% (18/24) of the batches infected with stumpy parasites, and as few as 10 stumpy parasites produced mature infections in immune adult flies. We discuss that, although Schuster et al. have demonstrated the intrinsic capacity of slender form trypanosomes to infect young and naive tsetse flies, highlighting the remarkable plasticity and adaptability of these protists, this phenomenon is unlikely to significantly contribute to the epidemiology of African trypanosomiases. According to both experimental and field observations, stumpy forms appear to be the most adapted forms for African trypanosome transmission from the mammalian host to the tsetse fly vector in natural conditions.

Introduction

Protist parasites of the Trypanosoma brucei group cause Human African Trypanosomiasis (HAT), or sleeping sickness in humans, and nagana in cattle (2). They are transmitted by the blood feeding tsetse fly following a long (at least 3 weeks) and complex (at least 9 distinct stages) cyclical development (review in (3)). In the mammalian host’s blood circulation, proliferating slender trypanosomes differentiate into cell cycle-arrested stumpy cells upon quorum sensing when they reach high parasite densities (4-7). This differentiation is thought not only to regulate the parasite load in the reservoir host (8), but also to provide transmissible parasites adapted to pursue the life cycle in the vector host (9). Indeed, stumpy forms express several transcripts and proteins necessary to the next developmental stage in the insect, the procyclic form, including the Protein Associated to Differentiation 1 or PAD1 (5). For decades, arrest of the cell cycle and differentiation to the stumpy stage were presumed essential for the developmental progression of bloodstream trypanosomes to the insect stages.

Recently, Schuster et al. demonstrated that slender trypanosomes can also present some intrinsic characteristic of transmissible forms (PAD1 mRNAs and proteins) and are readily transmissible to both young male and female tsetse flies, where they can complete their complex life cycle (1), yet with a lower efficiency than stumpy forms (10). In their experimental conditions, a single slender parasite was sufficient for productive infection. However, these laboratory conditions are significantly different from what is encountered in the field. First, only young teneral flies (1-3 days post-eclosion) with an immature immune system were used. Second, in some experiments, chemical compounds altering the insect immune response (glutathione) and/or promoting the infection (N-acetyl-glucosamine) were added to the infective meal. To assess the importance of these parameters, we challenged the infectivity of slender bloodstream forms in adult tsetse flies, i.e. in conditions closer to the natural situation.

Material and methods

Strains, culture and in vitro differentiation

The AnTat 1.1E Paris pleomorphic strain of Trypanosoma brucei brucei was derived from a strain originally isolated from a bushbuck in Uganda in 1966 (11). Bloodstream form trypanosomes were cultivated in HMI-9 medium supplemented with 10% (v/v) FBS (12) at 37°C in 5% CO2. Proliferative slender cells were maintained at densities lower than 5.105 parasites/ml to prevent their natural quorum-sensing-dependent differentiation into stumpy forms. For in vitro slender to stumpy BSF differentiation, we used 8-pCPT-2′-O-Me-5′-AMP, a nucleotide analogue of 5’-AMP (BIOLOG Life Science Institute, Germany). Briefly, 2x106 pleomorphic AnTat 1.1E slender forms were incubated with 8-pCPT-2′-O-Me-5′-AMP (5 μM) for 48 h (13). Freshly differentiated stumpy forms and slender cells were then centrifuged at 1,400 × g for 10 minutes and resuspended at the appropriate densities in SDM-79 medium supplemented with 10% FBS. Cells were resuspended at either 103, 104 or 105 parasites / ml. Assuming individual bloodmeal volumes ranging between 10-100 μl, this would correspond to ingestions of 10-100, 100-1,000 or 1,000-10,000 parasites per condition.

Tsetse fly maintenance, infection and dissection

Glossina morsitans morsitans tsetse flies were maintained in Roubaud cages at 27°C and 70% hygrometry and fed through a silicone membrane with fresh mechanically defibrinated sheep blood (BCL, France). Adult (between 2 and 3 weeks after emergence) or teneral males (between 24h and 72h post-emergence) were allowed to ingest parasites through a silicone membrane and without any chemical supplement. A total of 3 to 5 independent biological replicates per condition were performed with batches of 50 flies per condition.

Flies were starved for at least 24 hours before being dissected blindly 28 to 31 days post-ingestion for isolation of all stages from the midgut and salivary glands. For recovery of all tsetse organs, after rapid isolation of the salivary glands in a first drop of phosphate buffer saline (PBS), whole tsetse alimentary tracts, from the distal part of the foregut to the Malpighian tubules, were dissected and arranged lengthways in another drop of PBS as previously described (14, 15). Isolated organs were then scrutinized under a microscope at 40x magnification by two independent readers and infection rates per organ were scored (16).

Immuno-fluorescence analysis (IFA)

Cultured parasites were washed in TDB and spread onto poly-L-lysine coated slides. For flash methanol fixation, slides were air-dried for 10 min, fixed in methanol at -20°C for 5 seconds and rehydrated for 20 min in PBS. For immunodetection of stumpy forms, slides were incubated for 1 h at 37°C with a rabbit polyclonal anti-PAD1 antibody (kindly provided by Keith Matthews, University of Edinburgh) (5) diluted at 1:300 in PBS containing 0.1% Bovine Serum Albumin (BSA). After 3 consecutive 5 min washes in PBS, a species and subclass-specific secondary antibody coupled to the Alexa 488 fluorochrome (Jackson ImmunoResearch) diluted at 1:1000 in PBS containing 0.1% BSA was applied for 1 h at 37°C. After washing in PBS, slides were finally stained with 4’,6-diamidino-2-phenylindole (DAPI, 1 μg/ml) for visualization of kinetoplast and nuclear DNA content and mounted under coverslips with ProLong antifade reagent (Invitrogen), as previously described (14). Slides were observed under an epifluorescence DMI4000 microscope (Leica) with a 100x oil objective (NA 1.4) to assess the proportion of PAD1-positive cells in the infective meals (n > 100 cells / condition).

Added in the revised version: Statistical analysis

Added in the revised version: Infections rates were compared by a two-sided ANOVA at 95% confidence with Prism V10.0.3 (GraphPad). For all conditions, MG infection rate differences between adult and teneral flies were statistically significant at p<0.001 for all conditions, excepted for 10 ST at p<0.02.

Results

Pleomorphic T. b. brucei bloodstream forms were either maintained in culture at a density lower than 5.105 parasites / ml to prevent quorum-sensing-induced differentiation and obtain only slender forms or induced for differentiation with a 5’-AMP nucleotide analogue to obtain mostly stumpy forms. The expression of PAD1 at the cell surface was assessed by immunofluorescence analysis prior to each experimental infection: no PAD1 expression was detected in the slender group, whereas an average of 63% (52% to 71%, n = 12 replicates) of the induced cells were expressing PAD1. Batches of 50 teneral (< 72h) or adult (2-3 weeks) male tsetse flies were fed in parallel with either slender or stumpy forms at densities corresponding to individual ingestion of about 10, 100 or 1,000 parasites per bloodmeal. In total, 1,384 flies from 12 distinct experimental infections were dissected about 4 weeks (28 to 31 days) after parasite ingestion. Infection rates in midguts and salivary glands were quantified and plotted for each condition (Figure 1 and Table S1).

Stumpy forms are the predominant transmissible forms.

Comparison of (A) midgut and (B) salivary gland infection rates in teneral (T) Vs. adult (A) tsetse flies (batches of 50 flies) infected with 10 to 100 (circles, 3 independent experiments), 100 to 1,000 (squares, 4 independent experiments) or 1,000 to 10,000 (triangles, 5 independent experiments) parasites in the slender (SL in blue) or stumpy (ST in red) forms.

After ingestion of slender forms, infected flies were observed in only 1 batch out of 24. This occurred in not yet fully immunocompetent teneral flies and with the highest number of ingested parasites (1,000 to 10,000 parasites). In contrast, midgut and salivary glands infected flies were observed in 75% (18/24) and 62.5% (15/24) of the batches infected with stumpy parasites, respectively. As few as 10 stumpy parasites produced mature infections in immunocompetent adult flies and the infection rates were similar whatever the amounts of stumpy forms ingested. However, in more susceptible non-immune teneral flies, the infection rates were increasing with the number of stumpy forms ingested.

Discussion

The recent findings of Schuster et al. (1) have opened a debate on the traditional view of the trypanosome life cycle where slender trypanosomes are considered as non-competent for cyclical development in the insect vector (17). The authors proposed that their observations could provide a solution to a long-lasting paradox, namely the successful transmission of parasites in chronic infections, despite low parasitemia.

Earlier version: However, Schuster et al. performed all their experimental infections in laboratory conditions that have been optimized for maximum transmission efficiency, with the use of teneral flies, and for some experiments, the addition of N-acetylglucosamine and glutathione. Teneral flies remain unfed up to 3-5 days after emergence from their puparium (1) and are known to be significantly more susceptible to trypanosome infection as they are immunologically immature (weak immune system and, leaky peritrophic matrix) (18-21). However, in the wild, tsetse flies remain teneral for 1 to 4 days maximum (before they take their first blood meal), although they subsequently live as adult flies for up to 9 months (22), which attenuates the possible epidemiological impact of teneral flies on trypanosome transmission. In addition, all infective meals were supplemented with 60 mM N-acetylglucosamine, an inhibitor of tsetse midgut lectins enhancing trypanosome infection (23). For infections with monomorphic parasites, the addition of 12.5 mM glutathione, an antioxidant that reduces the midgut environment, protects trypanosomes from cell death induced by reactive oxygen species (24). These chemical compounds inhibit the (already weak) immune response in teneral flies, and substantially enhance the chances for trypanosomes to develop in the insect vector.

Revised version: However, Schuster et al. performed all their experimental infections in laboratory conditions that have been optimized for maximum transmission efficiency, with the use of teneral flies, and for some experiments, the addition of N-acetylglucosamine and glutathione. Teneral flies remain unfed up to 3-5 days after emergence from their puparium (1) and are known to be significantly more susceptible to trypanosome infection as they are thought to be immunologically immature (weak immune system and, leaky peritrophic matrix) (18-21). However, in the wild, tsetse flies remain teneral for 1 to 4 days maximum (before they take their first blood meal), although they subsequently live as adult flies for up to 9 months (22), which may attenuate the possible epidemiological impact of teneral flies on trypanosome transmission. In addition, all infective meals were supplemented with 60 mM N-acetylglucosamine, an inhibitor of tsetse midgut lectins enhancing trypanosome infection (23). For infections with monomorphic parasites, the addition of 12.5 mM glutathione, an antioxidant that reduces the midgut environment, protects trypanosomes from cell death induced by reactive oxygen species (24). These chemical compounds inhibit the (already weak) immune response in teneral flies, and substantially enhance the chances for trypanosomes to develop in the insect vector.

Conditions are far less favourable in the field; hence we investigated and compared the infection potential of slender and stumpy forms in adult and teneral flies without the addition of chemicals. We observed that slender forms were not infective to adult tsetse flies and only at densities higher than 105 parasites / ml. However, in endemic areas, especially in Western Africa, parasitemia in confirmed cases are usually very low (< 104 parasites / ml in Guinea for instance) and it is necessary to concentrate parasites in blood prior to microscopic examination to increase sensitivity of parasitological diagnosis (2). Hence, the possible transmission of a few slender trypanosomes from the blood of individuals with a chronic infection is unlikely to explain the maintenance of the parasite circulation in tsetse populations.

Earlier version: By contrast, we observed that as few as 10 stumpy parasites are enough to produce mature infections in both teneral and adult flies, already with a significant efficiency. In patients with low parasitemia, the quorum-sensing-triggered differentiation of slender to stumpy forms could be compatible with, or even more adapted to, extravascular forms present in some tissues and organs, with a limited dilution of the parasites and parasite factors remaining concentrated locally. Indeed, extravascular PAD1-positive trypanosomes were detected in high numbers at least in adipose tissues (25) and in the dermis (26) of experimentally infected mice. The presence of PAD1-positive extravascular trypanosomes was also assessed in the skin of confirmed gambiense HAT cases and unconfirmed seropositive individuals in endemic areas (27) (and unpublished data). This suggests that stumpy trypanosomes accessible to tsetse flies are likely more abundant than previously estimated in individuals with low parasitemia. Schuster et al. have unequivocally demonstrated the intrinsic capacity of slender form trypanosomes to infect young and naive tsetse flies, highlighting the remarkable plasticity and adaptability of these protists. The fine understanding of the underlying cellular mechanisms and / or transient adaptations involved in this process remains an exciting challenge. However, this event is unlikely to contribute to the epidemiology of African trypanosomiases in natural settings. According to both experimental and field observations, stumpy forms appear to be the most adapted forms for African trypanosome transmission from the mammalian host to the tsetse fly vector in natural conditions.

Revised version: By contrast, we observed that as few as 10 stumpy parasites are enough to produce mature infections in both teneral and adult flies, already with a significant efficiency. In patients with low parasitemia, the quorum-sensing-triggered differentiation of slender to stumpy forms could be compatible with, or even more adapted to, extravascular forms present in some tissues and organs, with a limited dilution of the parasites and parasite factors remaining concentrated locally. Indeed, extravascular PAD1-positive trypanosomes were detected in high numbers at least in adipose tissues (25) and in the dermis (26) of experimentally infected mice. The presence of PAD1-positive extravascular trypanosomes was also assessed in the skin of confirmed gambiense HAT cases and unconfirmed seropositive individuals in endemic areas (27) (and unpublished data). This suggests that stumpy trypanosomes accessible to tsetse flies are likely more abundant than previously estimated in individuals with low parasitemia. Schuster et al. have demonstrated the intrinsic capacity of slender form trypanosomes to infect young and naive tsetse flies, highlighting the remarkable plasticity and adaptability of these protists. The fine understanding of the underlying cellular mechanisms and / or transient adaptations involved in this process remains an exciting challenge. However, this event is unlikely to contribute to the epidemiology of African trypanosomiases in natural settings. According to both experimental and field observations, stumpy forms appear to be the most adapted forms for African trypanosome transmission from the mammalian host to the tsetse fly vector in natural conditions.

Supporting information